---
title: NLTK
description: ""
---

## Insert exercise title here

```yaml
type: NormalExercise
key: 108cd7d469
xp: 100
```

Let's start with dividing a document into sentences. Normally the first thing we do when we want to analyse a document is to break it down into sentences.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}
text = "Hello everyone. Hope you enjoy learning NLP. All the best."
# Try to extract the three sentences.
# You can use split method of strings
```

`@sample_code`
```{python}

```

`@solution`
```{python}
sentences = text.split('.')
```

`@sct`
```{python}

```
